

# Herramientas para an谩lisis estad铆stico

En este taller implementar谩s modelos de *machine learning* para resolver diversos problemas.

## Habilidades en pr谩ctica
Al desarrollar este taller podr谩s verificar tu progreso para:

**1.** Diferenciar los tipos de modelos de *machine learning* que ofrece Scikit Learn y sus casos de uso comunes.<br>
**2.** Construir, entrenar y analizar los resultados de modelos b谩sicos de *machine learning* en Scikit Learn.

## Instrucciones
En cada uno de los siguientes ejercicios deber谩s escribir el c贸digo solicitado estrictamente en las celdas indicadas para ello, teniendo en cuenta las siguientes recomendaciones:

* No crear, eliminar o modificar celdas de este Notebook (salvo lo que se te indique), pues puede verse afectado el proceso de calificaci贸n autom谩tica.
* La calificaci贸n se realiza de manera autom谩tica con datos diferentes a los proporcionados en este taller. Por consiguiente, tu c贸digo debe funcionar para diferentes instancias de cada uno de los ejercicios; una instancia hace referencia al valor de los par谩metros.
* La calificaci贸n de cada ejercicio depende del valor que retorne la funci贸n especificada en su enunciado. Por lo tanto, aunque implementes funciones adicionales, es escencial que utilices los nombres propuestos en los enunciados de los ejercicios para implementar la funci贸n definitiva.

Esta es una actividad calificada y, por lo tanto, debe ser resuelta individualmente.

## Ejercicios

Importa todos los paquetes que consideres necesarios para el desarollo de este taller en la siguiente celda.

```python
# Esta celda no es modificable

import numpy as np
import pandas as pd
from sklearn.model_selection import train_test_split
from sklearn.linear_model import LinearRegression
```

### Ejercicio 1

A continuaci贸n, encuentras un archivo de Excel con el cual podr谩s probar tu soluci贸n.

```python
archivo = pd.read_excel("Archivos/datosTaller.xlsx")
archivo
```

Output:
```
y  x1  x2  x3
0    9   6   3  22
1    1   8   1  35
2    5   6   2  41
3    3   4   2   3
4    9   5   1  18
5    9   7   4  18
6    6   8   5   9
7    4   7   5  30
8    2   6   3  16
9   10   7   5  47
10   3   6   4  50
11   7   5   4  11
12   6   3   4  42
13   8   5   4   1
```

Implementa una funci贸n llamada `dividir_datos` que reciba un `DataFrame` por par谩metro y retorne dos `DataFrame` contenidos en una tupla de la forma ($X$, $y$), donde $X$ es un `DataFrame` cuyas columnas son las variables independientes y $y$ es un `Series` con la variable de respuesta.

**Nota:** la variable de respuesta se encuentra en la primera columna de los datos.

*Esta funci贸n debe retornar una tupla de la forma ($X$, $y$) donde $X$ es un `DataFrame` y $y$ un `Series`*

<!-- grade_id: ej1-sol, points: 0, grade: false, solution: true -->
```python
# Implementa tu respuesta en esta celda
def dividir_datos(df):
    y = df.iloc[:, 0]   
    X = df.iloc[:, 1:]  
    return (X, y)
```

<!-- grade_id: ej1-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
import numpy.testing as npt
archivo = pd.read_excel("Archivos/datosTaller.xlsx")
archivo_test = pd.read_excel("Archivos/datosTallerTest.xlsx")

# Caso 1: no existe la funci贸n.
try:
    dividir_datos
    assert type(dividir_datos) == type(lambda:None)
except:
    raise NotImplementedError("No existe una funci贸n llamada dividir_datos.",)
    
# Caso 2: la funci贸n es interrumpida por errores durante su ejecuci贸n.
try:
    dividir_datos(archivo)
    dividir_datos(archivo_test)
except:
    raise RuntimeError("Tu funci贸n produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame como X.
assert type(dividir_datos(archivo)[0]) == pd.DataFrame, f"Tu funci贸n debe retornar una tupla de con un objeto en su primera posici贸n de tipo {type(archivo).__name__}."

# Caso 4: no retorna un Series como y.
assert type(dividir_datos(archivo)[1]) == pd.Series, f"Tu funci贸n debe retornar una tupla de con un objeto en su segunda posici贸n de tipo {type(pd.Series()).__name__}."

# Caso 5: respuesta y
try:
    npt.assert_approx_equal(dividir_datos(archivo)[1].mean(), 5.857142857142857)
    assert dividir_datos(archivo)[1][6] == 6
except AssertionError as e:
    e.args += ("Los datos de la variable de respuesta, y, son incorrectos.",)
    raise e

# Caso 6: respuesta X
try:
    npt.assert_approx_equal(dividir_datos(archivo)[0].mean()[1], 3.357142857142857)
    assert dividir_datos(archivo)[0].iloc[5,1] == 4
except AssertionError as e:
    e.args += ("Los datos de las variables explicativas, X, son incorrectos.",)
    raise e

# Caso 7: procedimiento incorrecto y
assert dividir_datos(archivo_test)[1].mean() == 27.65 and dividir_datos(archivo_test)[1][6] == 38, "Tu respuesta de la variable de respuesta, y, es incorrecta para una instancia diferente de los datos."

# Caso 8: procedimiento incorrecto X
assert dividir_datos(archivo_test)[0].mean()[1] == 294.9 and dividir_datos(archivo_test)[0].iloc[5,3] == 21, "Tu respuesta para X es incorrecta para una instancia diferente de los datos."

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
/var/folders/08/14s4v21s1hlgf09dhb1z76380000gn/T/ipykernel_74289/3598828985.py:38: FutureWarning: Series.__getitem__ treating keys as positions is deprecated. In a future version, integer keys will always be treated as labels (consistent with DataFrame behavior). To access a value by position, use `ser.iloc[pos]`
  npt.assert_approx_equal(dividir_datos(archivo)[0].mean()[1], 3.357142857142857)
/var/folders/08/14s4v21s1hlgf09dhb1z76380000gn/T/ipykernel_74289/3598828985.py:48: FutureWarning: Series.__getitem__ treating keys as positions is deprecated. In a future version, integer keys will always be treated as labels (consistent with DataFrame behavior). To access a value by position, use `ser.iloc[pos]`
  assert dividir_datos(archivo_test)[0].mean()[1] == 294.9 and dividir_datos(archivo_test)[0].iloc[5,3] == 21, "Tu respuesta para X es incorrecta para una instancia diferente de los datos."
```

### Ejercicio 2

A continuaci贸n, encuentras una tupla ($X$, $y$) con la cual podr谩s probar tu soluci贸n.

```python
datos = (pd.DataFrame({"x1":[6,  8,  6,  4, 5,  7,  8],
                       "x2":[3,  1,  2,  2, 1,  4,  5],
                       "x3":[22, 35, 41, 3, 18, 18, 9]}),
         pd.Series([9, 1, 5, 3, 9, 9, 6]))
datos
```

Output:
```
(   x1  x2  x3
 0   6   3  22
 1   8   1  35
 2   6   2  41
 3   4   2   3
 4   5   1  18
 5   7   4  18
 6   8   5   9,
 0    9
 1    1
 2    5
 3    3
 4    9
 5    9
 6    6
 dtype: int64)
```

Implementa una funci贸n llamada `crear_train_test` que reciba una tupla ($X$, $y$) por par谩metro y realice una partici贸n, 70% para entrenar al modelo y 30% para validarlo, con el fin de retornar una tupla de la forma (`X_train`, `X_test`, `y_train`, `y_test`). Usa el par谩metro `random_state = 0` para que tus respuestas coincidan con las nuestras.

Esta funci贸n debe retornar una tupla de la forma (`X_train`, `X_test`, `y_train`, `y_test`) donde `X_train` y `X_test` son de tipo `DataFrame` y `y_train` y `y_test` de tipo `Series`.

<!-- grade_id: ej2-sol, points: 0, grade: false, solution: true -->
```python
# Implementa tu respuesta en esta celda

def crear_train_test(datos):
    X, y = datos  
    X_train, X_test, y_train, y_test = train_test_split(
        X, y, test_size=0.3, random_state=0
    )
    return (X_train, X_test, y_train, y_test)
```

<!-- grade_id: ej2-test, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
datos = (pd.DataFrame({"x1":[6,  8,  6,  4, 5,  7,  8],
                       "x2":[3,  1,  2,  2, 1,  4,  5],
                       "x3":[22, 35, 41, 3, 18, 18, 9]}),
         pd.Series([9, 1, 5, 3, 9, 9, 6]))
archivo = pd.read_excel("Archivos/datosTaller.xlsx")
archivo_test = pd.read_excel("Archivos/datosTallerTest.xlsx")
datos_test = (pd.DataFrame(archivo)[pd.DataFrame(archivo).columns[1:]], pd.DataFrame(archivo)[pd.DataFrame(archivo).columns[0]])

# Caso 1: no existe la funci贸n.
try:
    crear_train_test
    assert type(crear_train_test) == type(lambda:None)
except:
    raise NotImplementedError("No existe una funci贸n llamada crear_train_test.",)
    
# Caso 2: la funci贸n es interrumpida por errores durante su ejecuci贸n.
try:
    crear_train_test(datos)
    crear_train_test((pd.DataFrame(archivo)[pd.DataFrame(archivo).columns[1:]], pd.DataFrame(archivo)[pd.DataFrame(archivo).columns[0]]))
except:
    raise RuntimeError("Tu funci贸n produce un error al ejecutarse.")

# Caso 3: no retorna un DataFrame como X_train.
assert type(crear_train_test(datos)[0]) == pd.DataFrame and type(crear_train_test(datos)[1]) == pd.DataFrame, f"Tu funci贸n debe retornar una tupla cuyos primeros dos objetos sean de tipo {pd.DataFrame.__name__}."

# Caso 4: no retorna un Series como y.
assert type(crear_train_test(datos)[2]) == pd.Series and type(crear_train_test(datos)[3]) == pd.Series, f"Tu funci贸n debe retornar un objeto de tipo '{pd.Series.__name__}' como y_train y y_test."

# Caso 5: respuesta X_train
assert crear_train_test(datos)[0].mean()[2] == 15.25 and crear_train_test(datos)[0].iloc[1,1] == 3, "Los valores de X_train son incorrectos."

# Caso 6: respuesta X_test
try:
    npt.assert_approx_equal(crear_train_test(datos)[1].mean()[2], 28.333333333333332)
    assert crear_train_test(datos)[1].iloc[1,1] == 2
except AssertionError as e:
    e.args += ("Los valores de X_test son incorrectos.",)
    raise e

# Caso 7: respuesta y_train
assert crear_train_test(datos)[2].mean() == 7.5 and crear_train_test(datos)[2].iloc[3] == 9, "La variable de respuesta y_train es incorrecta."

# Caso 8: respuesta y_test
assert crear_train_test(datos)[3].mean() == 4.0 and crear_train_test(datos)[3].iloc[2]==1, "La variable de respuesta y_test es incorrecta."

# Caso 9: procedimiento incorrecto X_train
try:
    npt.assert_approx_equal(crear_train_test(datos_test)[0].mean()[2], 27.555555555555557)
    assert crear_train_test(datos_test)[0].iloc[1,1] == 5
except AssertionError as e:
    e.args += ("Tu respuesta para X_train es incorrecta para una instancia diferente de los datos.",)
    raise e

# Caso 10: procedimiento incorrecto y_train
try:
    npt.assert_approx_equal(crear_train_test(datos_test)[2].mean(), 5.888888888888889)
    assert crear_train_test(datos_test)[2].iloc[3] == 4
except AssertionError as e:
    e.args += ("Tu respuesta de la variable de respuesta y_train es incorrecta para una instancia diferente de los datos.",)
    raise e

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
/var/folders/08/14s4v21s1hlgf09dhb1z76380000gn/T/ipykernel_74289/3579580357.py:33: FutureWarning: Series.__getitem__ treating keys as positions is deprecated. In a future version, integer keys will always be treated as labels (consistent with DataFrame behavior). To access a value by position, use `ser.iloc[pos]`
  assert crear_train_test(datos)[0].mean()[2] == 15.25 and crear_train_test(datos)[0].iloc[1,1] == 3, "Los valores de X_train son incorrectos."
/var/folders/08/14s4v21s1hlgf09dhb1z76380000gn/T/ipykernel_74289/3579580357.py:37: FutureWarning: Series.__getitem__ treating keys as positions is deprecated. In a future version, integer keys will always be treated as labels (consistent with DataFrame behavior). To access a value by position, use `ser.iloc[pos]`
  npt.assert_approx_equal(crear_train_test(datos)[1].mean()[2], 28.333333333333332)
/var/folders/08/14s4v21s1hlgf09dhb1z76380000gn/T/ipykernel_74289/3579580357.py:51: Futur
```

### Ejercicio 3

A continuaci贸n, encuentras una tupla de la forma (`X_train`, `X_test`, `y_train`, `y_test`) con la cual podr谩s probar tu soluci贸n.

```python
datos = (pd.DataFrame({"x1":[6,  8,  6,  4,  5,  7,  8],
                       "x2":[3,  1,  2,  2,  1,  4,  5],
                       "x3":[22, 35, 41, 3,  18, 18, 9]}),
         pd.DataFrame({"x1":[7,  6,  5,  3,  5],
                       "x2":[5,  4,  4,  4,  4],
                       "x3":[47, 50, 11, 42, 1]}),
         pd.Series([9,  1, 5, 3, 9, 9, 6]),
         pd.Series([10, 3, 7, 6, 8]))
datos
```

Output:
```
(   x1  x2  x3
 0   6   3  22
 1   8   1  35
 2   6   2  41
 3   4   2   3
 4   5   1  18
 5   7   4  18
 6   8   5   9,
    x1  x2  x3
 0   7   5  47
 1   6   4  50
 2   5   4  11
 3   3   4  42
 4   5   4   1,
 0    9
 1    1
 2    5
 3    3
 4    9
 5    9
 6    6
 dtype: int64,
 0    10
 1     3
 2     7
 3     6
 4     8
 dtype: int64)
```

Implementa una funci贸n llamada `crear_modelo` que reciba una tupla de la forma (`X_train`, `X_test`, `y_train`, `y_test`) por par谩metro y cree un modelo de regresi贸n lineal con el fin de explicar la variable de inter茅s. Finalmente, retorna el modelo creado luego de entrenarlo con los datos recibidos por par谩metro.

Esta funci贸n debe retornar un modelo de la librer铆a `sklearn`.

<!-- grade_id: ej3-sol, points: 0, grade: false, solution: true -->
```python
# Implementa tu respuesta en esta celda

def crear_modelo(datos):
    X_train, X_test, y_train, y_test = datos 
    modelo = LinearRegression()               
    modelo.fit(X_train, y_train)              
    return modelo
```

<!-- grade_id: cell-34296e646a810f8d, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
datos = (pd.DataFrame({"x1":[6,  8,  6,  4,  5,  7,  8],
                       "x2":[3,  1,  2,  2,  1,  4,  5],
                       "x3":[22, 35, 41, 3,  18, 18, 9]}),
         pd.DataFrame({"x1":[7,  6,  5,  3,  5],
                       "x2":[5,  4,  4,  4,  4],
                       "x3":[47, 50, 11, 42, 1]}),
         pd.Series([9,  1, 5, 3, 9, 9, 6]),
         pd.Series([10, 3, 7, 6, 8]))
archivo = pd.read_excel("Archivos/datosTaller.xlsx")
archivo_test = pd.read_excel("Archivos/datosTallerTest.xlsx")
datos_test = (pd.DataFrame(archivo_test)[pd.DataFrame(archivo_test).columns[1:]], pd.DataFrame(archivo_test)[pd.DataFrame(archivo_test).columns[0]])
datos_test = train_test_split(datos_test[0], datos_test[1], random_state=0)

# Caso 1: no existe la funci贸n.
try:
    crear_modelo
    assert type(crear_modelo) == type(lambda:None)
except:
    raise NotImplementedError("No existe una funci贸n llamada crear_modelo.",)
    
# Caso 2: la funci贸n es interrumpida por errores durante su ejecuci贸n.
try:
    crear_modelo(datos)
    #crear_modelo(archivo_test)
except:
    raise RuntimeError("Tu funci贸n produce un error al ejecutarse.")

# Caso 3: no retorna LinearRegression.
assert type(crear_modelo(datos)) == LinearRegression, f"Tu funci贸n debe retornar un objeto de tipo '{LinearRegression.__name__}'."

# Caso 4: respuesta coef_
try:
    npt.assert_approx_equal(crear_modelo(datos).coef_[0], -1.2195141060248855)
except AssertionError as e:
    e.args += ("Los coeficientes del modelo son incorrectos, revisa la implementaci贸n de la regresi贸n.",)
    raise e

# Caso 5: respuesta intercept_
try:
    npt.assert_approx_equal(crear_modelo(datos).intercept_, 7.996285638991168)
except AssertionError as e:
    e.args += ("El intercepto del modelo es incorrecto, revisa la implementaci贸n de la regresi贸n.",)
    raise e

# Caso 6: procedimiento incorrecto coef_
try:
    npt.assert_approx_equal(crear_modelo(datos_test).coef_[0], -1.953081558461383)
except AssertionError as e:
    e.args += ("Los coeficientes del modelo son incorrectos para una instancia diferente de los datos.",)
    raise e

# Caso 7: procedimiento incorrecto
try:
    npt.assert_approx_equal(crear_modelo(datos_test).intercept_, 37.752378744745855)
except AssertionError as e:
    e.args += ("El intercepto del modelo es incorrecto para una instancia diferente de los datos.",)
    raise e

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
```

### Ejercicio 4

A continuaci贸n, encuentras un modelo de regresi贸n lineal entrenado y una tupla de la forma (`X_train`, `X_test`, `y_train`, `y_test`) con la cual podr谩s probar tu soluci贸n. Para esto se usar谩 la librer铆a `pickle` la cual permite almacenar y cargar objetos de Python como modelos o estructuras de datos.

```python
import pickle

modelo = pickle.load(open("Archivos/modeloTaller", 'rb'))

datos = (pd.DataFrame({"x1":[6,  8,  6,  4,  5,  7,  8],
                       "x2":[3,  1,  2,  2,  1,  4,  5],
                       "x3":[22, 35, 41, 3,  18, 18, 9]}),
         pd.DataFrame({"x1":[7,  6,  5,  3,  5],
                       "x2":[5,  4,  4,  4,  4],
                       "x3":[47, 50, 11, 42, 1]}),
         pd.Series([9,  1, 5, 3, 9, 9, 6]),
         pd.Series([10, 3, 7, 6, 8]))
```

Output:
```
/Library/Frameworks/Python.framework/Versions/3.13/lib/python3.13/site-packages/sklearn/base.py:442: InconsistentVersionWarning: Trying to unpickle estimator LinearRegression from version 0.22.2.post1 when using version 1.7.1. This might lead to breaking code or invalid results. Use at your own risk. For more info please refer to:
https://scikit-learn.org/stable/model_persistence.html#security-maintainability-limitations
  warnings.warn(
```

Implementa una funci贸n llamada `predecir` que reciba un modelo entrenado y una tupla de la forma (`X_train`, `X_test`, `y_train`, `y_test`) por par谩metro y prediga la variable de inter茅s a partir del conjunto de datos `X_test`. Finalmente, la funci贸n debe retornar el valor de la predicci贸n hecha. 

Esta funci贸n debe retornar un arreglo de `numpy` con las predicciones hechas.

<!-- grade_id: ej4-sol, points: 0, grade: false, solution: true -->
```python
import numpy.testing as npt   #  necesario


def predecir(modelo, datos):
    X_train, X_test, y_train, y_test = datos
    if not hasattr(modelo, "positive"):
        modelo.positive = False

    if hasattr(modelo, "feature_names_in_") and isinstance(X_test, pd.DataFrame):
        req = list(modelo.feature_names_in_)
        faltantes = [c for c in req if c not in X_test.columns]
        if faltantes:
            raise ValueError(f"Faltan columnas requeridas por el modelo: {faltantes}")
        X_test = X_test[req]

    pred = modelo.predict(X_test)
    return np.asarray(pred)
```

<!-- grade_id: cell-6c66f9c02d416bdb, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
modelo = pickle.load(open("Archivos/modeloTaller", 'rb'))
datos = (pd.DataFrame({"x1":[6,  8,  6,  4,  5,  7,  8],
                       "x2":[3,  1,  2,  2,  1,  4,  5],
                       "x3":[22, 35, 41, 3,  18, 18, 9]}),
         pd.DataFrame({"x1":[7,  6,  5,  3,  5],
                       "x2":[5,  4,  4,  4,  4],
                       "x3":[47, 50, 11, 42, 1]}),
         pd.Series([9,  1, 5, 3, 9, 9, 6]),
         pd.Series([10, 3, 7, 6, 8]))
archivo = pd.read_excel("Archivos/datosTaller.xlsx")
archivo_test = pd.read_excel("Archivos/datosTallerTest.xlsx")
datos_test = (pd.DataFrame(archivo)[pd.DataFrame(archivo).columns[1:]], pd.DataFrame(archivo)[pd.DataFrame(archivo).columns[0]])
datos_test = train_test_split(datos_test[0], datos_test[1], random_state=0)

# Caso 1: no existe la funci贸n.
try:
    predecir
    assert type(predecir) == type(lambda:None)
except:
    raise NotImplementedError("No existe una funci贸n llamada predecir.",)
    
# Caso 2: la funci贸n es interrumpida por errores durante su ejecuci贸n.
try:
    predecir(modelo, datos)
    #crear_modelo(archivo_test)
except:
    raise RuntimeError("Tu funci贸n produce un error al ejecutarse.")

# Caso 3: no retorna un arreglo.
assert type(predecir(modelo, datos)) == np.ndarray, f"Tu funci贸n debe retornar un objeto de tipo '{np.ndarray.__name__}'."

# Caso 4: respuesta len()
assert len(predecir(modelo, datos)) == 5, "La longitud del arreglo con los datos dados es incorrecta."

# Caso 5: respuesta incorrecta
try:
    npt.assert_approx_equal(predecir(modelo, datos).mean(), 10.627099878205998)
    npt.assert_approx_equal(predecir(modelo, datos)[3], 13.866931147214054)
    npt.assert_approx_equal(predecir(modelo, datos)[1], 10.803128182847708)
except AssertionError as e:
    e.args += ("Los valores de la soluci贸n son incorrectos, revisa la implementaci贸n de la funci贸n.",)
    raise e

# Caso 6: procedimiento incorrecto
try:
    npt.assert_approx_equal(predecir(modelo, datos_test).mean(), 6.888353715429375)
    npt.assert_approx_equal(predecir(modelo, datos_test)[3], 9.123287939544587)
    npt.assert_approx_equal(predecir(modelo, datos_test)[1], 6.917762337875083)
except AssertionError as e:
    e.args += ("Los valores de la soluci贸n son incorrectos para una instancia diferente de los datos.",)
    raise e

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
/Library/Frameworks/Python.framework/Versions/3.13/lib/python3.13/site-packages/sklearn/base.py:442: InconsistentVersionWarning: Trying to unpickle estimator LinearRegression from version 0.22.2.post1 when using version 1.7.1. This might lead to breaking code or invalid results. Use at your own risk. For more info please refer to:
https://scikit-learn.org/stable/model_persistence.html#security-maintainability-limitations
  warnings.warn(
/Library/Frameworks/Python.framework/Versions/3.13/lib/python3.13/site-packages/sklearn/utils/validation.py:2742: UserWarning: X has feature names, but LinearRegression was fitted without feature names
  warnings.warn(
/Library/Frameworks/Python.framework/Versions/3.13/lib/python3.13/site-packages/sklearn/utils/validation.py:2742: UserWarning: X has feature names, but LinearRegression was fitted without feature names
  warnings.warn(
/Library/Frameworks/Python.framework/Versions/3.13/lib/python3.13/
```

### Ejercicio 5

A continuaci贸n, encuentras en la variable `modelo` un modelo de regresi贸n lineal entrenado y las variable `X_test` y `y_test` con las cuales podr谩s probar tu soluci贸n. Para esto se usar谩 la librer铆a `pickle` la cual permite almacenar y cargar objetos de Python como modelos o estructuras de datos.

```python
import pickle

modelo = pickle.load(open("Archivos/modeloTaller", 'rb'))

X_test = pd.DataFrame({"x1":[7,  6,  5,  3,  5],
                       "x2":[5,  4,  4,  4,  4],
                       "x3":[47, 50, 11, 42, 1]})

y_test = pd.Series([10, 3, 7, 6, 8])
```

Output:
```
/Library/Frameworks/Python.framework/Versions/3.13/lib/python3.13/site-packages/sklearn/base.py:442: InconsistentVersionWarning: Trying to unpickle estimator LinearRegression from version 0.22.2.post1 when using version 1.7.1. This might lead to breaking code or invalid results. Use at your own risk. For more info please refer to:
https://scikit-learn.org/stable/model_persistence.html#security-maintainability-limitations
  warnings.warn(
```

Implementa una funci贸n llamada `calcular_precision` que reciba un modelo entrenado, y una tupla de la forma ($X_{test}$, $y_{test}$) por par谩metro y retorne el score $R^2$ del modelo.

Esta funci贸n debe retornar un valor n煤merico tipo float (`flt`).

<!-- grade_id: ej5-sol, points: 0, grade: false, solution: true -->
```python
# Implementa tu respuesta en esta celda

def calcular_precision(modelo, X_test, y_test):
    if not hasattr(modelo, "positive"):
        modelo.positive = False

    return np.float64(modelo.score(X_test, y_test))

```

<!-- grade_id: cell-2c839c39325ca2bc, points: 1, grade: true, solution: false -->
```python
## AUTO-CALIFICADOR

# Base variables
import pickle

modelo = pickle.load(open("Archivos/modeloTaller", 'rb'))
X_test = pd.DataFrame({"x1":[7,  6,  5,  3,  5],
                       "x2":[5,  4,  4,  4,  4],
                       "x3":[47, 50, 11, 42, 1]})
y_test = pd.Series([10, 3, 7, 6, 8])
archivo = pd.read_excel("Archivos/datosTaller.xlsx")
archivo_test = pd.read_excel("Archivos/datosTallerTest.xlsx")
datos_test = (pd.DataFrame(archivo)[pd.DataFrame(archivo).columns[1:]], pd.DataFrame(archivo)[pd.DataFrame(archivo).columns[0]])
datos_test = train_test_split(datos_test[0], datos_test[1], random_state=0)

# Caso 1: no existe la funci贸n.
try:
    calcular_precision
    assert type(calcular_precision) == type(lambda:None)
except:
    raise NotImplementedError("No existe una funci贸n llamada calcular_precision.",)
    
# Caso 2: la funci贸n es interrumpida por errores durante su ejecuci贸n.
try:
    calcular_precision(modelo, X_test, y_test)
    #crear_modelo(archivo_test)
except:
    raise RuntimeError("Tu funci贸n produce un error al ejecutarse.")

# Caso 3: no retorna un float.
assert type(calcular_precision(modelo, X_test, y_test)) == np.float64, f"Tu funci贸n debe retornar un objeto de tipo '{np.float64.__name__}'."

# Caso 4: respuesta incorrecta
try:
    npt.assert_approx_equal(calcular_precision(modelo, X_test, y_test), -3.7894053047960226)
except AssertionError as e:
    e.args += ("La respuesta de la precisi贸n es incorrecta.",)
    raise e

# Caso 5: procedimiento incorrecto
try:
    npt.assert_approx_equal(calcular_precision(modelo, datos_test[1], datos_test[3]), -0.7120120937366505)
except AssertionError as e:
    e.args += ("El valor de la precisi贸n es incorrecto para una instancia diferente de los datos.",)
    raise e

# Mensaje de felicitaciones
print("Felicidades, realizaste este ejercicio correctamente.")
```

Output:
```
Felicidades, realizaste este ejercicio correctamente.
/Library/Frameworks/Python.framework/Versions/3.13/lib/python3.13/site-packages/sklearn/base.py:442: InconsistentVersionWarning: Trying to unpickle estimator LinearRegression from version 0.22.2.post1 when using version 1.7.1. This might lead to breaking code or invalid results. Use at your own risk. For more info please refer to:
https://scikit-learn.org/stable/model_persistence.html#security-maintainability-limitations
  warnings.warn(
/Library/Frameworks/Python.framework/Versions/3.13/lib/python3.13/site-packages/sklearn/utils/validation.py:2742: UserWarning: X has feature names, but LinearRegression was fitted without feature names
  warnings.warn(
/Library/Frameworks/Python.framework/Versions/3.13/lib/python3.13/site-packages/sklearn/utils/validation.py:2742: UserWarning: X has feature names, but LinearRegression was fitted without feature names
  warnings.warn(
/Library/Frameworks/Python.framework/Versions/3.13/lib/python3.13/
```

## Cr茅ditos

__Autor(es)__: Ariadna Sof铆a De vila Bula, Camilo Falla Moreno, Juan David Reyes Jaimes, Alejandro Mantilla Redondo, Diego Alejandro Cely G贸mez
 
__Fecha 煤ltima actualizaci贸n__: 04/10/2021